## Домашнее задание №3 

*Задачи 6-7 заимствованы из [домашнего задания](http://nbviewer.jupyter.org/url/python.math-hse.info/static/assignments_release/dj2017/ps02/ps02.ipynb) по [курсу](http://math-info.hse.ru/s17/1) «Введение в программирование» (Магистерская программа «Журналистика данных», НИУ ВШЭ, 2017-18). *

За разные задачи можно получить разное число баллов. Если не указано обратное, задача весит 1 балл. Максимум за ДЗ можно набрать 6 баллов. Вы можете решить больше задач, чем требуется, чтобы потренироваться.

Для предварительной проверки задания нужно сделать следующее:

1. Скачать данный `ipynb`-файл на свой компьютер, открыть его в IPython Notebook/Jupyter.
2. Активировать тесты (см. ниже).
3. Вставить решение каждой задачи в ячейку для кода, следующую за его условием, там, где написан комментарий `# YOUR CODE HERE`. Отступ вашего кода должен составлять 4 пробела. Ничего не менять вокруг!
4. Запустить ячейку, в которую вы вставили код с решением. Ввести какие-то входные данные, проверить визуально правильность вывода.
5. Запустить следующую ячейку (в ней содержится тест). Если запуск ячейки с тестом не приводит к появлению ошибки (assertion), значит, всё в порядке, задача решена. Если приводит к появлению ошибки, значит, тест не пройден и нужно искать ошибку.

**Внимание!** Если в какой-то момент забыть ввести входные данные и перейти на следующую ячейку, есть риск, что Notebook перестанет откликаться. В этом случае надо перезапустить ядро: *Kernel → Restart*. При этом потеряются все значения переменных, но сам код останется.

Чтобы сдать ДЗ, его надо загрузить в [nbgr-x](http://python.math-hse.info) в виде `ipynb`-файла. Получить `ipynb`-файл можно, выбрав в Jupyter пункт меню *File → Download as... → IPython Notebook (.ipynb)*.

### Активировать тесты
Запустите следующие ячейку, чтобы иметь возможность запускать тесты. Эту операцию нужно проделывать каждый раз, когда вы перезапускаете ядро. Если какой-то из тестов говорит `NameError: name 'Tester' is not defined`, нужно запустить эту ячейку ещё раз.

```python
# Фабрика тестов для проверки программ, принимающих данные через input()

from collections import deque

class Tester(object):
    def __init__(self, inp):
        self.outputs = []
        self.inputs = deque(inp)
    def print(self, *args, sep = " ", end = "\n"):
        text = sep.join(map(str, args)) + end
        newlines = text.splitlines(keepends=True)
        if self.outputs and self.outputs[-1] and self.outputs[-1][-1] != "\n" and newlines:
            self.outputs[-1] += newlines[0]
            self.outputs.extend(newlines[1:])
        else:
            self.outputs.extend(newlines)
            
    def input(self, *args):
        assert self.inputs, "Вы пытаетесь считать больше элементов, чем предусмотрено условием"
        return self.inputs.popleft()
    def __enter__(self):
        global print
        global input
        print = self.print
        input = self.input
        return self.outputs
    def __exit__(self, *args):
        global print
        global input
        del print
        del input
```

### Задача 1

Напишите программу, которая запрашивает у пользователя с клавиатуры целое число, а выводит на экран список из двух элементов: целое число, предшествующее ему, и целое число, следующее за ним.

**Примеры:**

*Входные данные:*

    2

*Выходные данные:*

    [1, 3]
    
*Входные данные:*

    0

*Выходные данные:*

    [-1, 1]

```python
def prev_next():
    # YOUR CODE HERE 
    
prev_next()    
```

```python
test_data = [
    ("2", [1, 3]),
    ("0", [-1, 1]),
    ("-1", [-2, 0]),
    ("101", [100, 102]),
    ("-10", [-11, -9])
]

for inp, out in test_data:
    with Tester([inp]) as t:
        prev_next()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == str(out) + "\n", "Неверный ответ, была введена строка " + inp
```

### Задача 2

Вася, студент второго курса образовательной программы с многообещающим названием, подрабатывает репетитором по английскому языку. Дан список его доходов за неделю:

    income = [2500, 0, 5000, 2500, 3000, 2500, 0]

Напишите программу, которая запрашивает у пользователя с клавиатуры число дней с начала недели, а возвращает сумму, заработанную Васей за все эти дни.

**Примеры:**

*Входные данные:*

    4

*Выходные данные:*

    10000
    
*Входные данные:*

    2

*Выходные данные:*

    2500
    
*Входные данные:*

    0

*Выходные данные:*

    0
    
**Подсказка:** сумму элементов списка можно получить с помощью функции `sum()`:

    sum([4, 6, 0])
    10

```python
def earnings():
    income = [2500, 0, 5000, 2500, 3000, 2500, 0]
    
    # YOUR CODE HERE
    
earnings()
```

```python
test_data = [
    ("0", "0"),
    ("1", "2500"),
    ("2", "2500"),
    ("3", "7500"),
    ("4", "10000"),
    ("5", "13000"),
    ("6", "15500"),
    ("7", "15500")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        earnings()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 3

Дан список станций желтой ветки московского метро (упорядочены как на карте):

    line = ["Третьяковская", "Марксистская", "Площадь Ильича", "Авиамоторная", "Шоссе Энтузиастов",
    "Перово", "Новогиреево", "Новокосино"]

Давайте считать, что движение происходит от Третьяковской до Новокосино.

Напишите программу, которая запрашивает у пользователя название текущей станции метро (без лишнего пробела на конце) и выводит на экран сообщение вида

    Следующая станция: станция.

Считайте, что пользователь знает, какая станция конечная, и не будет запрашивать станцию, следующую после Новокосино.

**Пример:**

*Входные данные:*

    Введите текущую станцию: Площадь Ильича

*Выходные данные:*

    Следующая станция: Авиамоторная

**Подсказка:**

Чтобы вернуть индекс определенного элемента в списке, нужно воспользоваться методом `.index()`. Пример:

    names = ["Anna", "Ella", "Max"]
    names.index("Ella")
    1

<!-- grade_id: p2, points: 0, grade: false, solution: true -->
```python
def metro():
    line = ["Третьяковская", "Марксистская", "Площадь Ильича", "Авиамоторная", "Шоссе Энтузиастов",
    "Перово", "Новогиреево", "Новокосино"]
    
    # YOUR CODE HERE

metro()
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("Третьяковская", "Следующая станция: Марксистская"),
    ("Марксистская", "Следующая станция: Площадь Ильича"),
    ("Площадь Ильича", "Следующая станция: Авиамоторная"),
    ("Авиамоторная", "Следующая станция: Шоссе Энтузиастов"),
    ("Шоссе Энтузиастов","Следующая станция: Перово"),
    ("Перово", "Следующая станция: Новогиреево"),
    ("Новогиреево", "Следующая станция: Новокосино")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        metro()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 4

Напишите программу, которая принимает на вход список неотрицательных чисел `N` и выводит на экран список квадратных корней из этих чисел (конечно, вычисление квадратных корней должно выполняться не вручную, хотя числа простые).

**Пример:**

*Входные данные:*

    L = [9, 16, 25, 36, 100, 81]

*Выходные данные:*

    [3.0, 4.0, 5.0, 6.0, 10.0, 9.0]

```python
def my_sqrt():
    L = [9, 16, 25, 36, 100, 81]
    
    # YOUR CODE HERE
    
my_sqrt()
```

```python
test_data = [([9, 16, 25, 36, 100, 81], "[3.0, 4.0, 5.0, 6.0, 10.0, 9.0]")]
for inp, out in test_data:
    with Tester([inp]) as t:
        my_sqrt()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert str(t[0]) == out+"\n", "Неверный ответ, была введена строка " + str(inp)
```

### Задача 5

Напишите программу, которая принимает на вход список слов такого вида:

    words = ["Speak ","to", "me ", "of", "Florence" ,"And ", "of", "the", "Renaissance"]

и выводит на экран список

    words_clean = ["speak", "to", "me", "of", "florence", "and", "of", "the", "renaissance"]

Другими словами, программа убирает пробелы в словах и приводит все слова к нижнему регистру.

**Подсказка:** используйте методы для строк `.strip()` и `.lower()`. Примеры:

    s = " ABC "
    s = s.strip()  # убирает лишние пробелы по краям
    s
    "ABC"
    
    s = a.lower()  # приводит все к нижнему регистру
    s
    "abc"

```python
def text_normalize():
    words = ["Speak ","to", "me ", "of", "Florence" ,"And ", "of", "the", "Renaissance"]
    
    # YOUR CODE HERE
    
text_normalize()
```

```python
test_data = [(["Speak ","to", "me ", "of", "Florence" ,"And ", "of", "the", "Renaissance"],
              ["speak", "to", "me", "of", "florence", "and", "of", "the", "renaissance"])]
for inp, out in test_data:
    with Tester([inp]) as t:
        text_normalize()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert str(t[0]) == str(out)+"\n", "Неверный ответ, была введена строка " + str(inp)
```

### Задача 6 (2 балла)
Факториалом натурального числа $n$ называется произведение всех натуральных чисел от 1 до $n$ включительно. Факториал $n$ обозначается $n!$. Например,
$$
4! = 1 \times 2 \times 3 \times 4 = 24.
$$
Написать программу, запрашивающую натуральное число $n$ с клавиатуры, вычисляющую и печатающую $n!$. Использовать какие-либо библиотечные функции нельзя (то есть запрещена конструкция `import`).

**Пример.**

**Входные данные:**

    5

**Выходные данные:**

    120

<!-- grade_id: p8, points: 0, grade: false, solution: true -->
```python
def my_fact():
    # YOUR CODE HERE

my_fact()
```

<!-- grade_id: p8_test, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("1", 1),
    ("2", 2),
    ("3", 6),
    ("4", 24),
    ("5", 120),
    ("10", 3628800),
    ("20", 2432902008176640000)
]

for inp, out in test_data:
    with Tester([inp]) as t:
        my_fact()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), "Неверный ответ, были введены числа " + inp
print("Хорошо, давайте теперь что-нибудь посложнее решим!")
```

### Задача 7 (2 балла)
Числа Фибоначчии — это последовательность чисел, заданная следующим образом:

$$a_1 = 1,\quad a_2 = 1, \quad a_{k+1}=a_k+a_{k-1}$$

Ввести число `k` с клавиатуры и вывести $a_k$. Считать, что число `k` всегда больше одного. (То есть на вход никогда не будет подано число, меньшее или равное одного.)


**Пример.**

    
_Входные данные:_

    3

_Выходные данные:_

    2
  
_Входные данные:_

    5

_Выходные данные:_

    5
  
_Входные данные:_

    6

_Выходные данные:_

    8

<!-- grade_id: p9, points: 0, grade: false, solution: true -->
```python
def fib_k():
    # YOUR CODE HERE
    
fib_k()
```

<!-- grade_id: p9_test, points: 2, grade: true, solution: false -->
```python
test_data = [(33, 3524578),
 (87, 679891637638612258),
 (66, 27777890035288),
 (43, 433494437),
 (44, 701408733),
 (19, 4181),
 (98, 135301852344706746049),
 (65, 17167680177565),
 (32, 2178309),
 (11, 89)]
for inp, out in test_data:
    with Tester([str(inp)]) as t:
        fib_k()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), ("Неверный ответ, было "
                                          "введено число " + str(inp) + 
                                         " ожидалось " + str(out))
print("Ура, всё верно!")
```